# SIADS 502: Math Methods for Data Science 


# School of Information, University of Michigan


## Week 3: Probability/Bayes Rule

### Version 1.1

This assignment is designed to refresh probability skills as well as apply probability in the context of Bayes' Theorem. Please read the directions carefully, as we want to avoid submissions that are marked incorrect due to formatting mistakes. 

Plese enter your name =""

## 1. Probability

<strong>1.1</strong> \[1 pt\] A college is looking to form a team of 8 to participate in a coding compeition. The team needs at least 2 females and 2 males. If there are 4 males and 7 females who want to join the competition, how many different ways can the team be formed? Please store your answer as an integer (e.g. ANS1 = 1) into the variable <strong>ANS1</strong>.

<!-- grade_id: cell-ee2bd7ab2781aa26, points: 0, grade: false, solution: true -->
```python
import math
import numpy as np
import numbers

ANS1=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2380314f88f56d22, points: 1, grade: true, solution: false -->
```python
assert type(ANS1) == int, "Problem 1.1, testing ANS1, type of value stored in variable does not match the expected type. Expecting Integer"
```

## 2. Conditional Probability

At a school, 40% of the students speak Spanish, 25% of the students speak German, and 15% of the students speak both of the langugues. Now, randomly pick a student. 

<strong>2.1</strong> \[1 pt\] Given that they speak Spanish, what is the probability that they also speak German? Store your answer in the variable <strong>ANS21</strong>.  For grading (on this and future problems) store the answer as accurately as possible using standard python data types.  For example, if the answer is 1/3, set ANS2 = 1/3 instead of .3333.

<!-- grade_id: cell-8566435eee91457e, points: 0, grade: false, solution: true -->
```python
ANS21 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-557a697b7229071a, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS21, numbers.Number), "Problem 2.1, testing ANS21, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>2.2</strong> \[1 pt\] Given that they do not speak German, what is the probability that they speak Spanish? Store answer in variable <strong>ANS22</strong>.  

<!-- grade_id: cell-d8a6247556748dcc, points: 0, grade: false, solution: true -->
```python
ANS22=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0c3c06d832ee2f42, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS22, numbers.Number), "Problem 2.2, testing ANS22, type of value stored in variable does not match the expected type. Expecting Float."
```

## 3. Independent Probability 

Tom and Lily are trying to solve a problem. The probability of Tom successfully solving the problem is 0.4 and the probability of Lily successfully solving the problem is 0.5. They do not influence each other when they are solving the problem. For each solution, store the value as accurately as you can (ex. ANS = 2/9 instead of .2222).

<strong>3.1</strong> \[1 pt\] What is the probability that both of them sucessfully solve the problem? Please store your answer to the variable <strong>ANS31</strong>.

<!-- grade_id: cell-a89dfc965132afad, points: 0, grade: false, solution: true -->
```python
ANS31 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d4bf277ef89e5a1, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS31, numbers.Number), "Problem 3.1, testing ANS31, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>3.2</strong> \[1 pt\] What is the probability that only one of them successfully solves the problem? Please store your answer to the variable <strong>ANS32</strong>. 

<!-- grade_id: cell-e55df99dc649e767, points: 0, grade: false, solution: true -->
```python
ANS32=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-04bf980d4b26a9f0, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS32, numbers.Number), "Problem 3.2, testing ANS32, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>3.3</strong> \[1 pt\] Given that **only** one of them successfully solved the problem, what is the probability that the person is Tom?
Please store your answer to the variable <strong>ANS33</strong>. 

<!-- grade_id: cell-9cd919f33acab2ef, points: 0, grade: false, solution: true -->
```python
ANS33 = None 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9beab75a7dab793d, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS33, numbers.Number), "Problem 3.3, testing ANS33, type of value stored in variable does not match the expected type. Expecting Float."
```

## 4. Bayes Rule

Suppose that you are an art dealer for Van Gogh paintings. Whenever someone brings in a painting that they claim is from Van Gogh, you send in the painting for pigment analysis. Pigment analysis successfully identifies 90% of Van Gogh paintings and falsely identifies 20% of forgeries as being from Van Gogh. In your experience, 99% of the paintings that people bring to you are forgeries.

<strong>4.1</strong> \[1 pt\] Use Bayes theorem to calculate the chance that a painting is from Van Gogh if the pigment analysis finds that it is from Van Gogh. Store your answer in the variable <strong>ANS41</strong> as accurately as possible.

<!-- grade_id: cell-f68142d6a4ebf595, points: 0, grade: false, solution: true -->
```python
ANS41 = None
# YOUR CODE HERE
raise NotImplementedError()
#print(type(ANS41))
```

<!-- grade_id: cell-c32efd59d0bc94f7, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS41, numbers.Number), "Problem 4.1, testing ANS41, type of value stored in variable does not match the expected type. Expecting Float."
```

## 5. Expectation and Variance

In a state park, a lost person arrives at a fork with 3 roads: <br>
• The first road leads back to the same point after 1 hour of walking.<br>
• The second road leads back to the same point after 2 hours of walking.<br>
• The third road leads to the parking lot after 5 hours of walking. <br>
There are no signs on the roads.

<strong>5.1</strong> \[2 pt\]  Assuming that the lost person always chooses each road with equal probability, what is the expected amount of time until the tourist arrives at the parking lot?  Store your answer in the variable <strong>ANS51</strong> as accurately as possible.

<!-- grade_id: cell-f68142d6a4ebf596, points: 0, grade: false, solution: true -->
```python
ANS51 = None
# YOUR CODE HERE
raise NotImplementedError()
#print(type(ANS41))
```

<!-- grade_id: cell-c32efd59d0bc94f3, points: 2, grade: true, solution: false -->
```python
assert isinstance(ANS51, numbers.Number), "Problem 5.1, testing ANS51, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>5.2</strong> \[2 pt\]  What is the expectation and variance of the sum of the roll of two dice?   Store your answer as a tuple in the variable <strong>ANS52</strong> as accurately as possible.

<!-- grade_id: cell-bbd9380cfee320a3, points: 0, grade: false, solution: true -->
```python
ANS52 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-402ceeaf0297d261, points: 2, grade: true, solution: false -->
```python
assert len(ANS52) == 2,  "Problem 5.2, testing ANS52, type of value stored in variable does not match the expected type. Tuple." 
assert isinstance(ANS52[0], numbers.Number) and isinstance(ANS52[1], numbers.Number) , "Problem 5.2, testing ANS52, type of values stored in tuples do not match the expected type. Expecting Float."
```

## 6. Monte Carlo Method



<strong>6.1</strong> \[2 pt\]  You receive one of five distinct coupons uniformily at random each day.  Use the Monte Carlo Method to estimate to within .01 days, the expected number of days until you collect two copies of each of the five coupons.  Record your answer in <strong>ANS61</strong>.

<!-- grade_id: cell-ceb4705829e49a2c, points: 0, grade: false, solution: true -->
```python
ANS61 = None
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-24a0cb9f1d006a9b, points: 2, grade: true, solution: false -->
```python
assert isinstance(ANS61, numbers.Number), "Problem 6.1, testing ANS61, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>6.2</strong> \[2 pt\]  There are 30 coins.  While 29 of them are fair, 1 of them flips heads with probability 60%.  You flip all 30 coins 100 times.  Use the Monte Carlo method to compute within .1% the probability that the unfair coin is in set of 10 coins that land heads most often?  (Count the event, that the biased coin is tied with the 10th most common fair coin).   Record your answer in <strong>ANS62</strong>.

Hint: use np.random.binomial to speed up simulation.  A binomial variable with parameters n and p is the number of heads resulting from flipping n coins, where each has probability p of landing heads. 

<!-- grade_id: cell-49d189d155b7bffe, points: 0, grade: false, solution: true -->
```python
ANS62 = None
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-691ed4f14f4a3360, points: 2, grade: true, solution: false -->
```python
assert isinstance(ANS62, numbers.Number), "Problem 6.2, testing ANS62, type of value stored in variable does not match the expected type. Expecting Float."
```